evaluation_interval: 1m

rule_files:
  - prometheus.quota_alerts.yaml

tests:
  - interval: 1m
    input_series:

      # namespace does not reach its quota limits
      - series: 'kube_resourcequota{type="hard", namespace="not-reaching-limits", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '2x15'
      - series: 'kube_resourcequota{type="used", namespace="not-reaching-limits", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '0x15'

      # namespace reaches its quota limits
      - series: 'kube_resourcequota{type="hard", namespace="reaching-limits", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '2x15'
      - series: 'kube_resourcequota{type="used", namespace="reaching-limits", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '2x15'

      # namespace exceeds its quota limits
      - series: 'kube_resourcequota{type="hard", namespace="exceeding-limits", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '2x15'
      - series: 'kube_resourcequota{type="used", namespace="exceeding-limits", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '5x15'
    alert_rule_test:
      - eval_time: 15m
        alertname: QuotaExceeded
        exp_alerts:
          - exp_labels:
              severity: warning
              namespace: exceeding-limits
              resource: example-resource
              resourcequota: test-quota
              source_cluster: cluster01
            exp_annotations:
              summary: >-
                Resource example-resource in namespace exceeding-limits exceeded quota
                test-quota
              description: >-
                Resource example-resource in namespace exceeding-limits on cluster
                cluster01 exceeded quota test-quota.
              runbook_url: https://gitlab.cee.redhat.com/rhtap/docs/sop/-/blob/main/o11y/alert-rule-QuotaExceeded.md

  - interval: 1m
    input_series:

      # Namespace exceeds its quota limits but since the name ends with 'env' it's ignored.
      - series: 'kube_resourcequota{type="hard", namespace="test-env", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '2x15'
      - series: 'kube_resourcequota{type="used", namespace="test-env", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '5x15'

      # Namespace exceeds its quota limits but since it's named 'tekton-ci' it's ignored.
      - series: 'kube_resourcequota{type="hard", namespace="tekton-ci", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '2x15'
      - series: 'kube_resourcequota{type="used", namespace="tekton-ci", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '5x15'

      # Namespace exceeds its quota limits but since it's named 'build-definitions-e2e' it's ignored.
      - series: 'kube_resourcequota{type="hard", namespace="build-definitions-e2e", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '2x15'
      - series: 'kube_resourcequota{type="used", namespace="build-definitions-e2e", resourcequota="test-quota", resource="example-resource", source_cluster="cluster01"}'
        values: '5x15'

    alert_rule_test:
      - eval_time: 15m
        alertname: QuotaExceeded
